# MIDAS Python Workshop, Afternoon, 2023
## Part 1: Python basics

### Problem 1
Write a function that takes two integers as arguments. Return the first argument if it is a positive number. Otherwise, return the sum of both arguments.

<!-- grade_id: cell-db7233a2d95c6406, points: 0, grade: false, solution: true -->
```python
def positive_or_sum(x, y):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-d525ecbed5402db1, points: 1, grade: true, solution: false -->
```python
assert positive_or_sum(42, 16) == 42
assert positive_or_sum(33, 13) == 33
assert positive_or_sum(-33, 13) == -20
```

### Problem 2
Write a function named `sum_naturals_while` that takes one argument `n` and use a `while` loop to compute the sum of the first $n$ natural numbers (starting with $1$).

<!-- grade_id: cell-9045223e53391c1c, points: 0, grade: false, solution: true -->
```python
def sum_naturals_while(n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-78c04715a9273632, points: 0, grade: true, solution: false -->
```python
assert sum_naturals_while(1) == 1
assert sum_naturals_while(2) == 3
assert sum_naturals_while(3) == 6
```

### Problem 3: Fun with Strings
A palindrome is a word or phrase that reads the same backwards and forwards. Write a function called `is_palindrome`, which takes a string as its only argument, and returns a Boolean. Your function should ignore spaces and capitalization in assessing whether or not a string is a palindrome, so that "tacocat" and "T A C O cat" are both considered palindromes.

<!-- grade_id: cell-7ae0586d5accba08, points: 0, grade: false, solution: true -->
```python
def is_palindrome(string):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-50df4deb8bcd2e47, points: 0, grade: true, solution: false -->
```python
assert type(is_palindrome('level')) == bool
assert is_palindrome('level') == True
assert is_palindrome('WAS it a car or a cat I saw') == True
assert is_palindrome(' ') == True
```

### Problem 4: Fun with Lists
1. <b>(2 points)</b> Write a function `list_reverse` that takes a list as an argument and returns that list, reversed. Your function should perform error checking to ensure that the input is a list. 

<!-- grade_id: cell-12f09bb2be5b3a31, points: 0, grade: false, solution: true -->
```python
def list_reverse(l):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-57b598c882f76f13, points: 0, grade: true, solution: false -->
```python
assert list_reverse (['this should be the same list']) == ['this should be the same list']
assert type(list_reverse([1,2,3,9])) == list
assert list_reverse([1, 2, 3]) == [3, 2, 1]
```

### Problem 5
Write a function `is_sorted` that takes a sequence `seq` as its only argument and returns True if the sequence is sorted in ascending order and returns False otherwise.

<!-- grade_id: cell-17e459df08cb0306, points: 0, grade: false, solution: true -->
```python
def is_sorted(seq):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-211d6372e5b2d3e4, points: 0, grade: true, solution: false -->
```python
assert is_sorted('abbde') == True
assert is_sorted([1, 2, 3, 3, 5]) == True
```

### Problem 6: Fun with Tuples
Write a function called `my_sum` that takes any number of numeric (ints and floats) arguments, and returns the sum of its arguments.

<!-- grade_id: cell-90314c57dc46cc64, points: 0, grade: false, solution: true -->
```python
def my_sum(*vector):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-7b6659c43e10dc6d, points: 0, grade: true, solution: false -->
```python
assert my_sum(3.2,4.1,-3) == 4.3
assert my_sum(3,4) == 7
```

### Problem 7: Function with Dictionaries
Write a function that concatenates two dictionaries to create a new one. You can assume the two input dictionaries have no keys in common.

<!-- grade_id: cell-c38834f58aa2c195, points: 0, grade: false, solution: true -->
```python
def dict_concat(dict1, dict2):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-78d03728839b5f2c, points: 0, grade: true, solution: false -->
```python
assert dict_concat({"rice": 4}, {"beans": 7}) == {"rice": 4, "beans": 7}
```

### Problem 8: Rotating tuples
Write a function called `rotate_tuple` that takes two arguments: a tuple and an integer, in that order. Letting $n$ be the integer given in the input, your function should return a tuple of the same length as the input tuple, but with its entries “rotated” by n. If n is positive, this should mean to “push forward” all the entries of the input tuple by n entries, with entries that “go off the end” of the tuple being wrapped around to the beginning, so that the $i^{th}$ entry of the input tuple becomes the $(i + n)^{th}$ entry of the output, wrapping around to the beginning of the tuple if this index goes off the end. If n is negative, then this corresponds to rotating the entries in the other direction, with entries of the input tuple being “pushed backward”.

<!-- grade_id: cell-12a83966d1802c7f, points: 0, grade: false, solution: true -->
```python
def rotate_tuple(t, n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-575be7dc46c02658, points: 0, grade: true, solution: false -->
```python
assert type(rotate_tuple((1,2,3,4), 5)) == tuple
assert rotate_tuple((1,2,3,4,5,6,7),2) == (6,7,1,2,3,4,5)
assert rotate_tuple((1,2,3,4,5,6,7),-2) == (3,4,5,6,7,1,2)
```

## Part 2: Numpy
### Problem 9
Convert a numpy array of Celsius measurements to Fahrenheit.

<!-- grade_id: cell-0a45e37e73575c3a, points: 0, grade: true, solution: true -->
```python
import numpy as np

def convert_to_fahrenheit(celsius_temps):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-5f36d4ed13034b2d, points: 0, grade: true, solution: false -->
```python
assert np.all(np.abs(convert_to_fahrenheit(np.zeros(4)) - 32.) < 1e-8)
assert np.all(np.abs(convert_to_fahrenheit(np.ones((10, 10)) * 100) - 212.) < 1e-8)
```

### Problem 10
Use numpy to create a function that adds a border filled with zeros around an existing array.

*Hint*: Use numpy's `pad` function. See https://numpy.org/doc/stable/reference/generated/numpy.pad.html

<!-- grade_id: cell-6f84a164211b781b, points: 0, grade: false, solution: true -->
```python
def add_border(x):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-fc67e1622f9dde62, points: 0, grade: true, solution: false -->
```python
x = np.ones((1, 1))
claimed = add_border(x)
correct = np.zeros((3, 3))
correct[1, 1] = 1.0
assert claimed.shape == correct.shape
assert np.all(np.abs(claimed - correct) < 1e-8)
```

### Problem 11
Suppose \begin{equation} 
f(x) = \begin{cases} \frac{ \sqrt{4-x^2} }{ 2\pi }
			&\mbox{ if } -2 \le x \le 2 \\
		0 &\mbox{ otherwise.} \end{cases}
\end{equation}

Define a function `wigner_density` that takes a single number (integer or float) as its input and returns a float as its output, given by the value of $f$ evaluated at the input. Your function should operate equally well on Python ints/floats and on __numpy__ ints/floats, and you should be able to accomplish this without checking the type of the input. Use the `numpy.sqrt` function for the square root, not the Python `math.sqrt` function.

<!-- grade_id: cell-d64e9a02985d3ca8, points: 0, grade: false, solution: true -->
```python
def wigner_density(x):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-f51fe2e26153aee3, points: 0, grade: true, solution: false -->
```python
assert wigner_density(2) == 0
assert wigner_density(1) == np.sqrt(3)/(2.0*np.pi)
```

### Problem 12
The density of a multivariate normal distribution is
$f(x) = (2\pi)^{-k/2}\mathrm{det}(\Sigma)^{-1/2}e^{-\frac{1}{2}(x - \mu)^\top \Sigma^{-1} (x - \mu)}$, where $x$ is the data, $k$ is the dimension of the data, $\mu$ is the mean, and $\Sigma$ is the covariance.

Write a function that computes this density using only numpy (and built-in Python commands);
do not use `scipy.stats.multivariate_normal`.

<!-- grade_id: cell-2f7e04935c5a8af9, points: 0, grade: false, solution: true -->
```python
def mvn_pdf(x, mean, cov):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-1ed23f7e13b77de7, points: 0, grade: true, solution: false -->
```python
from scipy.stats import multivariate_normal
import numpy as np

x = np.zeros(2)
mean = np.zeros(2)
cov = np.eye(2)
claimed = mvn_pdf(x, mean, cov)
rv = multivariate_normal(mean=mean, cov=cov)
correct = rv.pdf(x)
assert (claimed - correct) < 1e-8

x = np.ones(2)
mean = np.ones(2) * .7
mean[1] += .3
cov = np.eye(2) + 0.5
cov[0, 0] += 3
claimed = mvn_pdf(x, mean, cov)
rv = multivariate_normal(mean=mean, cov=cov)
correct = rv.pdf(x)
assert (claimed - correct) < 1e-8
```

## Part 3: Pandas
### Problem 13: Filtering and computing a mean
Write a function that returns the average rainfall on days with wind speed above 6 mph.

<!-- grade_id: cell-5e2bfa6d8f003e12, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b9e85016e94b2f74, points: 0, grade: false, solution: true -->
```python
def windy_rain():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-b494169726bf125a, points: 0, grade: true, solution: false -->
```python
assert abs(windy_rain() - 0.0412698) < 1e-6
```

### Problem 14: Filtering and computing a sum
Write a function that take a month as an argument and returns the total rainfall during the month.

<!-- grade_id: cell-41be4af1650f2589, points: 0, grade: false, solution: true -->
```python
def monthly_rain(month):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-41a23767c5b3538c, points: 0, grade: true, solution: false -->
```python
assert abs(monthly_rain("aug") - 10.8) < 1e-6
assert abs(monthly_rain("mar") - 0.2) < 1e-6
```

### Problem 15: Loading Pandas data frames
Download the white wine dataset from [https://archive.ics.uci.edu/ml/machine-learning-databases/wine-quality/winequality-white.csv](https://archive.ics.uci.edu/ml/machine-learning-databases/wine-quality/winequality-white.csv), and load it into a Pandas dataframe called `df`.

<!-- grade_id: cell-f263f757e0de4beb, points: 0, grade: false, solution: true -->
```python
from urllib import request
from os import path
import pandas as pd

url = "https://archive.ics.uci.edu/ml/machine-learning-databases/wine-quality/"
fn = "winequality-white.csv"
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-68523472e0373632, points: 0, grade: true, solution: false -->
```python
assert abs(df["citric acid"].mean() - 0.3341915) < 1e-8
```

### Problem 16: pipe
Write a function that takes a user-defined function as an argument.
It uses `pipe` to apply the function to each wine pH, and returns the mean of the values returned by the user-defined function.

<!-- grade_id: cell-cb307124b2001682, points: 0, grade: false, solution: true -->
```python
def mean_ph(user_fun):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-66349e8134132db2, points: 0, grade: true, solution: false -->
```python
assert abs(mean_ph(lambda x: x**2) - (df.pH ** 2).mean()) < 1e-8
```

### Problem 17: groupby
Use the Pandas `groupby` function to create a function that takes no arguments and returns the mean citric acid for wines with each quality level as a series indexed by quality.

<!-- grade_id: cell-f4f4d36f6804ca64, points: 0, grade: false, solution: true -->
```python
def quality_mean():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-68c8595f920e6379, points: 0, grade: true, solution: false -->
```python
assert abs(quality_mean()[9] - 0.386) < 1e-8
```